Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Rithvik Mundra"
COLLABORATORS = "Rithvik Mundra"
```

---

# Exercise 1 - Shell basics

Work through as much of the Software Carpentry [lesson on the Unix Shell](http://swcarpentry.github.io/shell-novice/) as you can.  Run through the Setup section just below, then open a terminal through Jupyter to run through the exercises.

After you have completed the first few sections of the tutorial, return to this notebook.  Execute all of the cells, and answer all of the questions.


## Setup - getting required files 

To get started, you'll need to have the required files in your directory.  Use `wget` to get them:

```python
!wget http://swcarpentry.github.io/shell-novice/data/shell-novice-data.zip
```

Output:
```
--2016-09-02 00:09:02--  http://swcarpentry.github.io/shell-novice/data/shell-novice-data.zip
Resolving swcarpentry.github.io (swcarpentry.github.io)... 151.101.32.133
Connecting to swcarpentry.github.io (swcarpentry.github.io)|151.101.32.133|:80... connected.
HTTP request sent, awaiting response... 200 OK
Length: 168355 (164K) [application/zip]
Saving to: ‘shell-novice-data.zip’

shell-novice-data.z 100%[===================>] 164.41K   545KB/s    in 0.3s    

2016-09-02 00:09:02 (545 KB/s) - ‘shell-novice-data.zip’ saved [168355/168355]
```

```python
!unzip -l shell-novice-data.zip
```

Output:
```
Archive:  shell-novice-data.zip
  Length      Date    Time    Name
---------  ---------- -----   ----
        0  2015-11-22 19:39   data-shell/
      199  2015-11-22 19:39   data-shell/.bash_profile
        0  2015-11-17 10:02   data-shell/creatures/
     1838  2015-11-17 10:02   data-shell/creatures/basilisk.dat
     1833  2015-11-17 10:02   data-shell/creatures/unicorn.dat
        0  2015-11-17 10:02   data-shell/data/
      283  2015-11-17 10:02   data-shell/data/amino-acids.txt
      136  2015-11-17 10:02   data-shell/data/animals.txt
        0  2015-11-17 10:02   data-shell/data/elements/
      270  2015-11-17 10:02   data-shell/data/elements/Ac.xml
      312  2015-11-17 10:02   data-shell/data/elements/Ag.xml
      314  2015-11-17 10:02   data-shell/data/elements/Al.xml
      325  2015-11-17 10:02   data-shell/data/elements/Am.xml
      307  2015-11-17 10:02   data-shell/data/elements/Ar.xml
      326  2015-11-17 10:02   data-shell/data/elements/As.xml
      261  2015-11-17 10:02
```

```python
!unzip shell-novice-data.zip
```

Output:
```
Archive:  shell-novice-data.zip
   creating: data-shell/
  inflating: data-shell/.bash_profile  
   creating: data-shell/creatures/
  inflating: data-shell/creatures/basilisk.dat  
  inflating: data-shell/creatures/unicorn.dat  
   creating: data-shell/data/
  inflating: data-shell/data/amino-acids.txt  
  inflating: data-shell/data/animals.txt  
   creating: data-shell/data/elements/
  inflating: data-shell/data/elements/Ac.xml  
  inflating: data-shell/data/elements/Ag.xml  
  inflating: data-shell/data/elements/Al.xml  
  inflating: data-shell/data/elements/Am.xml  
  inflating: data-shell/data/elements/Ar.xml  
  inflating: data-shell/data/elements/As.xml  
  inflating: data-shell/data/elements/At.xml  
  inflating: data-shell/data/elements/Au.xml  
  inflating: data-shell/data/elements/B.xml  
  inflating: data-shell/data/elements/Ba.xml  
  inflating: data-shell/data/elements/Be.xml  
  inflating: data-shell/data/elements/Bi.xml  
  inflating: data-shell/data/elements/Bk.xml  
  
```

*Note*: you only need to do this once per session while using Jupyter on datanotebook.org.  You can open a terminal now and work through the steps, and return to this notebook a little later, and the files will be available either way.  That's because you're working in the same local directory.

**However**!  If you download this file, close your Jupyter session for the night, then come back tomorrow and open up a new Jupyter session on the server again, you'll need to get those sample files again.  Just execute the cells above to do it.

Okay, let's get on with the exercise!

## Navigating Files and Directories

As you work through this section of the tutorial, complete the steps here as well, using the `!` shell escape command.  Execute each cell as you go.

These steps aren't exactly the same as what's in the tutorial, where the file layout is a little different and where they're not using a notebook like we are.  That's okay.  Just consider this practice.

```python
!whoami
```

Output:
```
rithvik
```

```python
!pwd
```

Output:
```
/home/rithvik/Downloads
```

```python
!ls -F
```

Output:
```
Anaconda3-4.1.1-Linux-x86_64.sh  shell-novice-data.zip
data-shell/			 test-file.txt
exercise-01.ipynb		 wk-01-a-Fall 2016-basic Python.ipynb
numbers.txt			 wk-01-b-Fall 2016-basic Python.ipynb
```

```python
!ls -F
```

Output:
```
Anaconda3-4.1.1-Linux-x86_64.sh  shell-novice-data.zip
data-shell/			 test-file.txt
exercise-01.ipynb		 wk-01-a-Fall 2016-basic Python.ipynb
numbers.txt			 wk-01-b-Fall 2016-basic Python.ipynb
```

```python
!ls -F data-shell/
```

Output:
```
creatures/  Desktop/	north-pacific-gyre/  pizza.cfg	writing/
data/	    molecules/	notes.txt	     solar.pdf
```

```python
!ls -aF
```

Output:
```
./				 numbers.txt
../				 shell-novice-data.zip
Anaconda3-4.1.1-Linux-x86_64.sh  test-file.txt
data-shell/			 wk-01-a-Fall 2016-basic Python.ipynb
exercise-01.ipynb		 wk-01-b-Fall 2016-basic Python.ipynb
.ipynb_checkpoints/
```

```python
!ls -af .
```

Output:
```
test-file.txt			 wk-01-b-Fall 2016-basic Python.ipynb
data-shell			 wk-01-a-Fall 2016-basic Python.ipynb
shell-novice-data.zip		 numbers.txt
Anaconda3-4.1.1-Linux-x86_64.sh  exercise-01.ipynb
.ipynb_checkpoints		 ..
.
```

What is the difference between the two previous cells, and what does the single dot mean?

<!-- grade_id: 156, points: 1, grade: true, solution: true -->
Answer:The command ls -aF shows all the files and directories sorted in the present working directory, clearly distinguishing the directories with a backslash("/") appended at the end whereas the command ls -af . shows all the files and directories unsorted and without any clear distiction between the directories and files.

```python
!ls -F ..
```

Output:
```
anaconda3/  Downloads/	      Pictures/  R/	      Templates/
Desktop/    examples.desktop  prog/	 rithvik/     Untitled.ipynb
Documents/  Music/	      Public/	 start.ipynb  Videos/
```

What do the double dots mean?

<!-- grade_id: 155, points: 1, grade: true, solution: true -->
Answer:The double dot represents the Directory one level above. Here my working directory is / home/rithvik/Downloads and hence when we use the command ls -F .. it will display all the files and directories present in the / home/rithvik.

```python
!ls data-shell/north-pacific-gyre/2012-07-03/
```

Output:
```
goodiff		NENE01736A.txt	NENE01843A.txt	NENE01978B.txt	NENE02040Z.txt
goostats	NENE01751A.txt	NENE01843B.txt	NENE02018B.txt	NENE02043A.txt
NENE01729A.txt	NENE01751B.txt	NENE01971Z.txt	NENE02040A.txt	NENE02043B.txt
NENE01729B.txt	NENE01812A.txt	NENE01978A.txt	NENE02040B.txt
```

## Working with Files and Directories

The following cells come from the next section of the tutorial.

```python
!ls -F
```

Output:
```
Anaconda3-4.1.1-Linux-x86_64.sh   shell-novice-data.zip
data-shell/			  test-file.txt
exercise-01 Rithvik Mundra.ipynb  wk-01-a-Fall 2016-basic Python.ipynb
numbers.txt			  wk-01-b-Fall 2016-basic Python.ipynb
```

```python
!mkdir thesis
```

<!-- grade_id: 121, points: 1, grade: true, solution: false -->
```python
import os
assert "thesis" in os.listdir()
```

```python
!ls -F
```

Output:
```
Anaconda3-4.1.1-Linux-x86_64.sh   test-file.txt
data-shell/			  thesis/
exercise-01 Rithvik Mundra.ipynb  wk-01-a-Fall 2016-basic Python.ipynb
numbers.txt			  wk-01-b-Fall 2016-basic Python.ipynb
shell-novice-data.zip
```

You can't use the nano editor here in Jupyter, so we'll use the `touch` command to create an empty file instead.

```python
!touch thesis/draft.txt
```

<!-- grade_id: 123, points: 1, grade: true, solution: false -->
```python
assert "draft.txt" in os.listdir("thesis")
```

```python
!ls -F thesis
```

Output:
```
draft.txt
```

Removing files and directories.

```python
!rm thesis/draft.txt
```

<!-- grade_id: 125, points: 1, grade: true, solution: false -->
```python
assert "draft.txt" not in os.listdir("thesis")
```

```python
!rm thesis
```

Output:
```
rm: cannot remove 'thesis': Is a directory
```

```python
!rmdir thesis
```

<!-- grade_id: 127, points: 1, grade: true, solution: false -->
```python
assert "thesis" not in os.listdir()
```

```python
!ls
```

Output:
```
Anaconda3-4.1.1-Linux-x86_64.sh   shell-novice-data.zip
data-shell			  test-file.txt
exercise-01 Rithvik Mundra.ipynb  wk-01-a-Fall 2016-basic Python.ipynb
numbers.txt			  wk-01-b-Fall 2016-basic Python.ipynb
```

Renaming and copying files.

```python
!touch draft.txt
```

<!-- grade_id: 129, points: 1, grade: true, solution: false -->
```python
assert "draft.txt" in os.listdir()
```

```python
!mv draft.txt quotes.txt
```

<!-- grade_id: 130, points: 1, grade: true, solution: false -->
```python
assert "quotes.txt" in os.listdir()
assert "draft.txt" not in os.listdir()
```

```python
!ls
```

Output:
```
Anaconda3-4.1.1-Linux-x86_64.sh   shell-novice-data.zip
data-shell			  test-file.txt
exercise-01 Rithvik Mundra.ipynb  wk-01-a-Fall 2016-basic Python.ipynb
numbers.txt			  wk-01-b-Fall 2016-basic Python.ipynb
quotes.txt
```

```python
!cp quotes.txt quotations.txt
```

<!-- grade_id: 134, points: 1, grade: true, solution: false -->
```python
assert "quotes.txt" in os.listdir()
assert "quotations.txt" in os.listdir()
```

```python
ls
```

Output:
```
Anaconda3-4.1.1-Linux-x86_64.sh   quotes.txt
[0m[01;34mdata-shell[0m/                       [01;31mshell-novice-data.zip[0m
exercise-01 Rithvik Mundra.ipynb  test-file.txt
numbers.txt                       wk-01-a-Fall 2016-basic Python.ipynb
quotations.txt                    wk-01-b-Fall 2016-basic Python.ipynb
```

```python

```